# CS211 Final Project 
Stanhope Nwosu and Clasby Chope 

Preforming workload queries and comparing methods of differential privacy on dataset of 911 calls in Montgomery Country PA

```python
# Load the data and libraries
import pandas as pd
import numpy as np
import random
from scipy import stats
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def laplace_mech_vec(vec, sensitivity, epsilon):
    return [v + np.random.laplace(loc=0, scale=sensitivity / epsilon) for v in vec]

def gaussian_mech(v, sensitivity, epsilon, delta):
    return v + np.random.normal(loc=0, scale=sensitivity * np.sqrt(2*np.log(1.25/delta)) / epsilon)

def gaussian_mech_vec(vec, sensitivity, epsilon, delta):
    return [v + np.random.normal(loc=0, scale=sensitivity * np.sqrt(2*np.log(1.25/delta)) / epsilon)
            for v in vec]

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

calls = pd.read_csv('https://raw.githubusercontent.com/stan0fHope/CS211Projecct/main/911.csv', low_memory=False)
calls = calls.dropna()
```

```python
calls['lat'].sort_values(ascending=True)

```

Output:
```
66737    30.333596
25529    32.387090
56937    39.745533
68123    39.958421
42440    39.958858
           ...    
13490    40.479853
4571     40.479853
54757    40.479853
4570     40.479853
12482    40.479853
Name: lat, Length: 75725, dtype: float64
```

## Range Queries

A *range query* counts the number of rows in the dataset which have a value lying in a given range. For example, "how many participants are between the ages of 21 and 33?" is a range query. A *workload* of range queries is just a list of range queries. The code below generates 100 random range queries over ages in the adult dataset.

```python
def range_query(df, col, a, b):
    return len(df[(df[col] >= a) & (df[col] < b)])

#can cahange for more ranges (18036, 19525) outlier 36107, 77316
random_zip_bounds = [random.randint(18000, 19000) for _ in range(100)]

#are float types tho so i cut 3 from top/bot
random_lat_bounds = [random.randint(25, 40) for _ in range(100)]

#sae for lat, but neg (-95, -74), make neg
random_lng_bounds = [random.randint(-100, -70) for _ in range(100)]

zip_workload = [(lb, random.randint(lb, 20000)) for lb in random_zip_bounds]
lat_workload = [(lb, random.randint(lb, 50)) for lb in random_lat_bounds]
lng_workload = [(lb, random.randint(lb, -65)) for lb in random_lng_bounds]

print('First 5 queries: ', zip_workload[:5])
real_zip = [range_query(calls, 'zip', lb, ub) for (lb, ub) in zip_workload]
real_lat = [range_query(calls, 'lat', lb, ub) for (lb, ub) in lat_workload]
real_lng = [range_query(calls, 'lng', lb, ub) for (lb, ub) in lng_workload]

```

Output:
```
First 5 queries:  [(18638, 18785), (18906, 19959), (18176, 19250), (18946, 19154), (18873, 18913)]
```

<!-- grade_id: cell-91ba0020d4ea347e, points: 0, grade: false, solution: true -->
```python
def workload_laplace(workload, epsilon, col):
    ans_list = []#apply range query
    for work in workload: #seq comp
        a, b = work
        rng = range_query(calls, col, a, b)
        noised = laplace_mech(rng, len(workload), epsilon)
        ans_list.append(noised)
    return ans_list

print('First 4 answers:', workload_laplace(zip_workload, 1.0, 'zip')[:4])
print('First 4 answers:', workload_laplace(lat_workload, 1.0, 'lat')[:4])
print('First 4 answers:', workload_laplace(lng_workload, 1.0, 'lng')[:4])
```

Output:
```
First 4 answers: [73320.53571668669, 43116.230298141614, 1188.2583679611173, 73688.05442616207]
First 4 answers: [74422.20961324422, 75705.27491674412, 75798.75822290042, 75723.75936368431]
First 4 answers: [75758.11690134865, 114.96128629832798, 75824.31939070918, 151.1356780022027]
```

<!-- grade_id: cell-3fdb7419fcfc0cea, points: 10, grade: true, solution: false -->
```python
errors = [abs(r_a - l_a) for (r_a, l_a) in zip(real_answers, workload_laplace(random_workload, 1.0))]
print('Average absolute error:', np.mean(errors))
assert np.mean(errors) > 50
assert np.mean(errors) < 200
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-59-b5819a56c38a> in <module>
----> 1 errors = [abs(r_a - l_a) for (r_a, l_a) in zip(real_answers, workload_laplace(random_workload, 1.0))]
      2 print('Average absolute error:', np.mean(errors))
      3 assert np.mean(errors) > 50
      4 assert np.mean(errors) < 200

NameError: name 'real_answers' is not defined
```

```python
##Our errors testing
zip_errors = [abs(r_a - l_a) for (r_a, l_a) in zip(real_zip, workload_laplace(zip_workload, 1.0, 'zip'))]
lat_errors = [abs(r_a - l_a) for (r_a, l_a) in zip(real_lat, workload_laplace(lat_workload, 1.0, 'lat'))]
lng_errors = [abs(r_a - l_a) for (r_a, l_a) in zip(real_lng, workload_laplace(lng_workload, 1.0, 'lng'))]

print('Average absolute error:', np.mean(zip_errors))
assert np.mean(zip_errors) > 50
assert np.mean(zip_errors) < 200

print('Average absolute error:', np.mean(lat_errors))
assert np.mean(lat_errors) > 50
assert np.mean(lat_errors) < 200

print('Average absolute error:', np.mean(lng_errors))
assert np.mean(lng_errors) > 50
assert np.mean(lng_errors) < 200

```

Output:
```
Average absolute error: 96.4690334016222
Average absolute error: 108.144619836204
Average absolute error: 110.78785978284162
```

## Question 2 (10 points)

Write code to answer a workload using `laplace_mech_vec` - the version of the Laplace mechanism for **vector-valued** queries. Your solution should *not* use sequential composition, and should have a total privacy cost of `epsilon`.

*Hint*: remember to use L1 global sensitivity.

<!-- grade_id: cell-e4e3cdb5562933b8, points: 0, grade: false, solution: true -->
```python
def workload_laplace_vec(workload, epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

print('First 4 answers:', workload_laplace_vec(random_workload, 1.0)[:4])
```

<!-- grade_id: cell-68af547658adcdcd, points: 10, grade: true, solution: false -->
```python
errors = [abs(r_a - l_a) for (r_a, l_a) in zip(real_answers, workload_laplace_vec(random_workload, 1.0))]
print('Average absolute error:', np.mean(errors))
assert np.mean(errors) > 50
assert np.mean(errors) < 200
```

## Question 3 (10 points)

In 2-5 sentences, answer the following:
- Did the two solutions differ in terms of their accuracy?
- How do they differ in terms of their use of composition properties of differential privacy?

<!-- grade_id: cell-0da21812c3f96147, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3 (10 points)

Write code to answer a workload using `gaussian_mech_vec` - the version of the Gaussian mechanism for vector-valued queries. Your solution should not use sequential composition, should satisfy $(\epsilon, \delta)$-differential privacy, and should have a total privacy cost of (`epsilon`, `delta`).

*Hint*: remember to use L2 sensitivity.

<!-- grade_id: cell-217b044c0cd1697e, points: 0, grade: false, solution: true -->
```python
def workload_gaussian_vec(workload, epsilon, delta):
    # YOUR CODE HERE
    raise NotImplementedError()

print('First 4 answers:', workload_gaussian_vec(random_workload, 1.0, 1e-5)[:4])
```

<!-- grade_id: cell-7c4c4d2ed403ea72, points: 10, grade: true, solution: false -->
```python
errors = [abs(r_a - l_a) for (r_a, l_a) in zip(real_answers, workload_gaussian_vec(random_workload, 1.0, 1e-5))]
print('Average absolute error:', np.mean(errors))
assert np.mean(errors) > 10
assert np.mean(errors) < 100
```

## Question 4 (10 points)

In 2-5 sentences, answer the following:
- Of your solutions in questions 1-3, which ones rely on *sequential composition*?
- Which solution offers the best accuracy?
- Why does this particular solution yield the best accuracy?

<!-- grade_id: cell-d505748348979709, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 5 (10 points)

Re-implement your solution to question 3 using *Rényi differential privacy*. Your solution should satisfy $(\alpha, \bar\epsilon)$-RDP.

*Hint*: see the "variants" chapter in the textbook.

<!-- grade_id: cell-50da3a173b751f3f, points: 0, grade: false, solution: true -->
```python
def workload_gaussian_vec_RDP(workload, alpha, epsilon_bar):
    # YOUR CODE HERE
    raise NotImplementedError()

print('First 4 answers:', workload_gaussian_vec(random_workload, 1.0, 1e-5)[:4])
```

<!-- grade_id: cell-d8eb625008d2985e, points: 10, grade: true, solution: false -->
```python
# TEST CASE
errors = [abs(r_a - l_a) for (r_a, l_a) in zip(real_answers, workload_gaussian_vec_RDP(random_workload, 5, 0.1))]
print('Average absolute error:', np.mean(errors))
assert np.mean(errors) > 10
assert np.mean(errors) < 100
```

## Question 6 (10 points)

Implement a function `convert_RDP_ED` to convert from the $(\alpha, \bar\epsilon)$ of Rényi differential privacy to the $(\epsilon, \delta)$ of approximate differential privacy. Your function should also take the desired value of $\delta$.

<!-- grade_id: cell-361a5a41cab1e483, points: 0, grade: false, solution: true -->
```python
def convert_RDP_ED(alpha, epsilon_bar, delta):
    # YOUR CODE HERE
    raise NotImplementedError()

convert_RDP_ED(5, 0.1, 1e-5)
```

<!-- grade_id: cell-98f137a5295dc1e9, points: 10, grade: true, solution: false -->
```python
# TEST CASE
assert convert_RDP_ED(5, 0.1, 1e-5) == 2.9782313662425572
assert convert_RDP_ED(40, 0.1, 1e-5) == 0.39520321705051864
assert convert_RDP_ED(500, 1.0, 1e-5) == 1.02307199491978
assert convert_RDP_ED(40, 1.0, 1e-5) == 1.2952032170505188
```

## Question 7 (10 points)

In 2-5 sentences, answer the following:
- Try various values for `alpha` and `epsilon_bar` in `convert_RDP_ED`. At what values do you observe an $(\epsilon, \delta)$ value around $(1.0, 10^{-5})$?
- Try these values for `alpha` and `epsilon_bar` in `workload_gaussian_vec_RDP`. How does the error compare to using `workload_gaussian_vec`?
- Is it useful to use Rényi differential privacy to answer workloads of range queries? Or is regular $(\epsilon, \delta)$-differential privacy just as good?

<!-- grade_id: cell-17fe437fb638503d, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE